Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

This first assignment is meant to determine your mastery the python essentials. 

It is set up so that if you can complete the assignment successfully, you're ready to begin the coding component of the course.

Don't worry if you get stuck! You can cover the essentials of coding in [python here](https://github.com/thisismetis/dbs_ml_curriculum/tree/master/curriculum/ml_intro/python-review)

---
## Part A (2 points)

Write a function that takes in a list of a class roster like the one above and returns a list of the first names of each student in the list. 


Make sure that the function does the following:
- If a first name appears twice in the roster, it should appear twice in what the function returns. 
- The output list should be sorted alphabetically.


<!-- grade_id: first_names, points: 0, grade: false, solution: true -->
```python
def get_first_names(class_list):
    """This function takes in a list of a class roster like the one above
        and returns a list of the first names of each 
        student in the list. If a first name appears twice in the roster, it 
        should  appear twice in what the function returns. 
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Write a test case for your function to make sure it works

```python
# test things out here
```

<!-- grade_id: first_names_test, points: 10, grade: true, solution: false -->
```python
spring_class_list = [
    'Timothy Butler', 'Heather Williams', 'Kimberly Harris', 'Lisa Miller',
    'Carrie Gonzalez', 'Nicholas Fitzgerald', 'Jessica Bird', 'Maria Steele',
    'Hannah Cruz', 'Heather Watkins'
]
spring_first_names = [
    'Carrie', 'Hannah', 'Heather', 'Heather', 'Jessica', 'Kimberly', 'Lisa',
    'Maria', 'Nicholas', 'Timothy'
]

assert get_first_names(spring_class_list) == spring_first_names
```

<!-- grade_id: first_names_edge, points: 5, grade: true, solution: false -->
```python
spring_class_list.append("heather billiams")
spring_first_names.insert(2, "Heather")

assert get_first_names(spring_class_list) == spring_first_names

spring_class_list.append("HEATHER Villiams")
spring_first_names.insert(2, "Heather")

assert get_first_names(spring_class_list) == spring_first_names
```

```python

```